<a href="https://colab.research.google.com/github/tmckim/materials-sp24-colab/blob/main/review/Q3_ReviewNotebook.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>


## Before you start - Save this notebook!

When you open a new Colab notebook from the WebCampus (like you hopefully did for this one), you cannot save changes. So it's  best to store the Colab notebook in your personal drive `"File > Save a copy in drive..."` **before** you do anything else.

The file will open in a new tab in your web browser, and it is automatically named something like: "**Copy of Q3_ReviewNotebook.ipynb**". You can rename this to just the title of the assignment "**Q3_ReviewNotebook.ipynb**". Make sure you do keep an informative name (like the name of the assignment) so that you know which files to submit back to WebCampus for grading! More instructions on this are at the end of the notebook.


**Where does the notebook get saved in Google Drive?**

By default, the notebook will be copied to a folder called “Colab Notebooks” at the root (home directory) of your Google Drive. If you use this for other courses or personal code notebooks, I recommend creating a folder for this course and then moving the assignments AFTER you have completed them. <br>

I also recommend you give the folder where you save your notebooks^ a different name than the folder we create below that will store the notebook resources you need each time you work through a course notebook. This includes any data files you will need, links to the images that appear in the notebook, and the files associated with the autograder for answer checking.<br>
You should select a name other than '**NS499-DataSci-course-materials**'. <br>
This folder gets overwritten with each assignment you work on in the course, so you should **NOT** store your notebooks in this folder that we use for course materials! <br><br>For example, you could create a folder called 'NS499-**notebooks**' or something along those lines.
___

### We will now do the setup steps as separate cells to help with issues finding files in google drive/colab. <br> If you restart colab, you must rerun all **5** steps in each of these cells!

```python
# Step 1
# Setup and add files needed to access gdrive
from google.colab import drive                                   # these lines mount your gdrive to access the files we import below
drive.mount('/content/gdrive', force_remount=True)
```

```python
# Step 2
# Change directory to the correct location in gdrive (modified way to do this from before)
import os
os.chdir('/content/gdrive/MyDrive/NS499-DataSci-course-materials/')
```

```python
# Step 3
# Remove the files that were previously there- we will replace with all the old + new ones for this assignment
!rm -r materials-sp24-colab
```

```python
# Step 4
# These lines clone (copy) all the files you will need from where I store the code+data for the course (github)
# Second part of the code copies the files to this location and folder in your own gdrive
!git clone https://github.com/tmckim/materials-sp24-colab '/content/gdrive/My Drive/NS499-DataSci-course-materials/materials-sp24-colab/'
```

```python
# Step 5
# Change directory into the folder where the resources for this assignment are stored in gdrive (modified way from before)
os.chdir('/content/gdrive/MyDrive/NS499-DataSci-course-materials/materials-sp24-colab/review/')
```

```python
# Import packages and other things needed
# Don't change this cell; Just run this cell
# If you restart colab, make sure to run this cell again after the first ones above^

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *
import d8error

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## Simulations & Defining Functions with Roulette

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner.

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="roulette_wheel.jpeg" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv')
wheel
```

### Betting on Red ###
If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- if the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

**Hint:** This is like what we practiced in class when we compared my roll of a die to yours with simulation. One outcome was positive and one was negative.

**Question 1.** Define a function `dollar_bet_on_red` that takes the name of a color and returns your gain in dollars if that color had won and you had placed a one-dollar bet on red. Remember that the gain can be negative. Make sure your function returns an integer.

*Note:* You can assume that the only colors that will be passed as arguments are red, black, and green. Your function doesn't have to check that.


```python
def dollar_bet_on_red(color):
```

```python
#@title Hidden Solution
def dollar_bet_on_red(color):
    if color == 'red':
        return 1
    else:
        return -1
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

**Question 2.** Now it's time to apply the function you just wrote. Your code should use the function `dollar_bet_on_red` on the correct column from the `wheel` table. The output array `red_winnings` should contain the gain in dollars if your pocket won and ou had bet one dollar on red.



```python
red_winnings = ...
```

```python
#@title Hidden Solution
red_winnings = wheel.apply(dollar_bet_on_red, 'Color')
```

```python
wheel = wheel.with_columns('Red Wins', red_winnings)
wheel
```

### Simulating 10 Bets on Red
Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

**Question 3.** Write a line of code that creates a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. Your table should have the same three column labels as in `wheel`.


```python
ten_bets = ...
```

```python
#@title Hidden Solution
ten_bets = wheel.sample(10)
```

```python
# let's find the sum of our winnings
total = sum(ten_bets.column('Red Wins'))
total
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

**Question 4.** Let's see what would happen if you made more bets. Define a function `net_gain_red` that takes the number of bets and returns the net gain in that number of one-dollar bets on red.

*Hint 1:* You should use your `wheel` table within your function, with similar code to what you wrote in the above question. <br>
*Hint 2:* you should return the sum of the column `Red Wins` from your new table created in the function below


```python
def net_gain_red(...):
  all_bets = ...
  return ...

```

```python
#@title Hidden Solution
def net_gain_red(num_bets):
    all_bets = wheel.sample(num_bets)
    return sum(all_bets.column('Red Wins'))
```

Run the cell below a few times to make sure that the results are similar to those you observed in the previous exercise.

```python
net_gain_red(10)
```

**Question 5.** Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. After the cell is run, `all_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red.

*Hint:* For this question, you can assume that `numpy` has been imported as `np`.


```python
num_bets = 200
repetitions = 10000
all_gains_red = make_array()


for ... in ...(...):
    new_net_gain = ...(...)
    all_gains_red = ...(..., ...)

len(all_gains_red)
```

```python
#@title Hidden Solution
num_bets = 200
repetitions = 10000
all_gains_red = make_array()

for i in np.arange(repetitions):
    new_net_gain = net_gain_red(num_bets)
    all_gains_red = np.append(all_gains_red, new_net_gain)

len(all_gains_red) # Do not change this line! Check that all_gains_red is length 10000.
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', all_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

**Extra Thought Question:** Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question.

```python
#@title Hidden Solution
loss_more_than_50 = True # SOLUTION
```

Tests for q1_1 (points: [2, 2])
```
>>> dollar_bet_on_red('red') == 1
True
```
```
>>> dollar_bet_on_red('green') == -1
True
```


Tests for q1_10 (points: [0, 0, 4])
```
>>> # Make sure histogram_statements is an array.
>>> import numpy as np
>>> type(histogram_statements) == np.ndarray
True
```
```
>>> # Make sure you have set histogram_statements to an array with at least 1 number
>>> 1 <= histogram_statements.item(0) <= 3
True
```
```
>>> histogram_statements.item(0) == 1 and histogram_statements.item(1) == 2 and histogram_statements.item(2) == 3
True
```


Tests for q1_2 (points: [0, 4])
```
>>> # Make sure your column names are correct
>>> wheel.labels[2] == "Winnings: Red"
True
```
```
>>> sum(wheel.column("Winnings: Red")) == -2
True
```


Tests for q1_3 (points: [0, 4])
```
>>> # If this test isn't passing, try running the cells from the top.
>>> set(["Pocket", "Color", "Winnings: Red"]) == set(ten_bets.labels)
True
```
```
>>> ten_bets.num_rows == 10
True
```


Tests for q1_4 (points: [0, 4])
```
>>> net_gain_red(10000) != net_gain_red(10000)
True
```
```
>>> -10000 <= net_gain_red(10000) <= 10000
True
```


Tests for q1_5 (points: [0, 4])
```
>>> len(all_gains_red) == 10000
True
```
```
>>> np.count_nonzero(all_gains_red <= 100) == 10000
True
```


Tests for q1_6 (points: [0, 4])
```
>>> type(loss_more_than_50) == bool
True
```
```
>>> loss_more_than_50 == True
True
```


Tests for q1_7 (points: [2, 1, 1])
```
>>> dollar_bet_on_split('5') == 17
True
```
```
>>> dollar_bet_on_split('6') == 17
True
```
```
>>> dollar_bet_on_split('10') == -1
True
```


Tests for q1_8 (points: [0, 4])
```
>>> # Make sure your column names are correct
>>> wheel.labels[3] == "Winnings: Split"
True
```
```
>>> sum(wheel.column("Winnings: Split")) == -2
True
```


Tests for q1_9 (points: [0, 4])
```
>>> len(all_gains_split) == 10000
True
```
```
>>> np.count_nonzero(all_gains_split >= -200) == 10000
True
```


Tests for q2_1 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < first_three_black < 1
True
```
```
>>> first_three_black
0.10628371482723427
```


Tests for q2_2 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < no_green < 1
True
```
```
>>> 0.582 < no_green < 0.583
True
```


Tests for q2_3 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < at_least_one_green < 1
True
```
```
>>> 0.417643346770 <= at_least_one_green <= 0.417643346771
True
```


Tests for q2_4 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < lone_winners < 1
True
```
```
>>> lone_winners
0.0011374153385028288
```


Tests for q2_5 (points: [0, 0, 4])
```
>>> # Make sure you are setting toys_option to an int
>>> type(toys_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= toys_option <= 3
True
```
```
>>> toys_option == 1
True
```


Tests for q2_6 (points: [0, 0, 4])
```
>>> # Make sure you are setting lottery_option to an int
>>> type(lottery_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= lottery_option <= 3
True
```
```
>>> lottery_option == 2
True
```


Tests for q2_7 (points: [0, 0, 4])
```
>>> # Make sure you are setting coin_option to an int
>>> type(coin_option) == int
True
```
```
>>> 1 <= coin_option <= 3
True
```
```
>>> coin_option == 1
True
```


Tests for q4_1 (points: [0, 0, 1, 3])
```
>>> # The array should have length 2
>>> len(deck_model_probabilities) == 2
True
```
```
>>> # The elements in the array should add up to 1.
>>> sum(deck_model_probabilities) == 1
True
```
```
>>> deck_model_probabilities.item(0) == 4/13
True
```
```
>>> deck_model_probabilities.item(1) == 9/13
True
```


Tests for q4_2 (points: [0, 0, 5])
```
>>> # Make sure you are setting statistic_choice to an int
>>> type(statistic_choice) == int
True
```
```
>>> # statistic_choice should be assigned to either 1, 2 or 3.
>>> 1 <= statistic_choice <= 3
True
```
```
>>> statistic_choice == 3
True
```


Tests for q4_3 (points: [5])
```
>>> # The statistic should be between 0 and 13 face cards for
>>> # a sample size of 13
>>> num_face = deck_simulation_and_statistic(13, deck_model_probabilities)
>>> 0 <= num_face <= 13
True
```


Tests for q4_4 (points: [2, 3])
```
>>> # There should be exactly as many elements in deck_statistics
>>> # as the number 'repetitions'
>>> len(deck_statistics) == repetitions
True
```
```
>>> # Each element of deck_statistics should be between 0
>>> # and 13 inclusive
>>> all([0 <= k <= 13 for k in deck_statistics])
True
```
